## Module 2

#### In this assignment, you will work on movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Load and read the 'imdb.xlsx' file. Read the 'imdb' sheet into a DataFrame, df.
"""

import pandas as pd

imdb = pd.read_csv("imdb.csv")

imdb

```

Output:
```
movie_title  director_id  country_id  \
0                         The Shawshank Redemption           34           1   
1                                    The Godfather           33           1   
2                                  The Dark Knight           16           1   
3                           The Godfather: Part II           33           1   
4    The Lord of the Rings: The Return of the King           83           1   
..                                             ...          ...         ...   
173                The Perks of Being a Wallflower          101           1   
174                                         Brazil          106           8   
175                                       Big Fish          108           1   
176                                   True Romance          112           1   
177                                     Persepolis          114           7   

    content_rating  title_year  imdb_score      gross  duration  
0                R      
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(df, sol.df)
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
""" Q2: 
Store the dimensions of the DataFrame as a tuple in a variable called 'shape' and print it.

Hint: A tuple is made up of comma separated values inside parenthesis.  e.g. (1, 2)
"""

# your code here
imdb.shape
```

Output:
```
(178, 8)
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(shape, sol.shape)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
""" Q3: 
Store the column titles and the types of data in variables named 'columns' and 'dtypes', then print them.
"""

# your code here
imdb.columns
```

Output:
```
Index(['movie_title', 'director_id', 'country_id', 'content_rating',
       'title_year', 'imdb_score', 'gross', 'duration'],
      dtype='object')
```

```python
imdb.dtypes
```

Output:
```
movie_title        object
director_id         int64
country_id          int64
content_rating     object
title_year          int64
imdb_score        float64
gross               int64
duration            int64
dtype: object
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(columns.all(), sol.columns.all())
assert_series_equal(dtypes, sol.dtypes)
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
""" Q4: 
Examine the first 10 rows of data; store them in a variable called first10
"""

# your code here
imdb.head(10)
```

Output:
```
movie_title  director_id  country_id  \
0                       The Shawshank Redemption           34           1   
1                                  The Godfather           33           1   
2                                The Dark Knight           16           1   
3                         The Godfather: Part II           33           1   
4  The Lord of the Rings: The Return of the King           83           1   
5                                   Pulp Fiction           85           1   
6                 The Good, the Bad and the Ugly           98           2   
7                               Schindler's List          103           1   
8                                      Inception           16           1   
9                                     Fight Club           22           1   

  content_rating  title_year  imdb_score      gross  duration  
0              R        1994         9.3   28341469       142  
1              R        1972         9.2  134821952       175
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(first10, sol.first10)
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
""" Q5: 
Examine the first 5 rows of data; store them in a variable called first5
"""

# your code here
imdb.head(5)
```

Output:
```
movie_title  director_id  country_id  \
0                       The Shawshank Redemption           34           1   
1                                  The Godfather           33           1   
2                                The Dark Knight           16           1   
3                         The Godfather: Part II           33           1   
4  The Lord of the Rings: The Return of the King           83           1   

  content_rating  title_year  imdb_score      gross  duration  
0              R        1994         9.3   28341469       142  
1              R        1972         9.2  134821952       175  
2          PG-13        2008         9.0  533316061       152  
3              R        1974         9.0   57300000       220  
4          PG-13        2003         8.9  377019252       192
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(first5, sol.first5)
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
""" Q6: 
Import the "directors" and "countries" sheets into their own DataFrames, df_directors and df_countries.
"""

# your code here
df_directors = pd.read_csv("directors.csv")
df_countries = pd.read_csv("countries.csv")
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_frame_equal(df_directors, sol.df_directors)
assert_frame_equal(df_countries, sol.df_countries)
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
""" Q7: 
Check the "directors" sheet
1. Count how many records there are based on the "id" column. (To get the number of records per "id", 
   use the value_counts method.) Store the result in a variable named count.
2. Remove the duplicates from the directors dataframe and store the result in a variable called df_directors_clean.
"""

# your code here

df_directors
```

Output:
```
id          director_name
0      1         Akira Kurosawa
1      2     Alejandro Amenabar
2      3  Alejandro G. Inarritu
3      3  Alejandro G. Inarritu
4      4       Alfred Hitchcock
..   ...                    ...
173  113         Victor Fleming
174  114      Vincent Paronnaud
175  115           Wes Anderson
176  116      Wolfgang Petersen
177  117            Woody Allen

[178 rows x 2 columns]
```

```python
df_directors.count()
```

Output:
```
id               178
director_name    178
dtype: int64
```

```python
df_directors.drop_duplicates(inplace=True)
```

```python
df_directors
```

Output:
```
id          director_name
0      1         Akira Kurosawa
1      2     Alejandro Amenabar
2      3  Alejandro G. Inarritu
4      4       Alfred Hitchcock
5      5         Andrew Stanton
..   ...                    ...
172  113         Victor Fleming
174  114      Vincent Paronnaud
175  115           Wes Anderson
176  116      Wolfgang Petersen
177  117            Woody Allen

[117 rows x 2 columns]
```

```python
df_directors_clean = df_directors
```

<!-- grade_id: q7_test, points: 2, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_series_equal(count, sol.count)
assert_frame_equal(df_directors_clean, sol.df_directors_clean)
```