## Module 3

#### In this assignment, you will continue working on the movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

We have loaded the data as "df" for you. Follow the instructions and finish the rest.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: init_user, points: 0, grade: true, solution: false -->
```python
# Loading the data
import pandas as pd

xls = pd.ExcelFile('imdb.xlsx')
df = xls.parse('imdb')
df_directors = xls.parse('directors')
df_countries = xls.parse('countries')

print("Data Loading Finished.")
```

Output:
```
Data Loading Finished.
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Join three Dataframes: df, df_directors, and df_countries with an inner join.
Store the joined DataFrames in df.
"""

# your code here
df

# After the join, the resulting Dataframe should have 12 columns.
#df.shape
```

Output:
```
movie_title  director_id  country_id  \
0                         The Shawshank Redemption           34           1   
1                                    The Godfather           33           1   
2                                  The Dark Knight           16           1   
3                           The Godfather: Part II           33           1   
4    The Lord of the Rings: The Return of the King           83           1   
..                                             ...          ...         ...   
173                The Perks of Being a Wallflower          101           1   
174                                         Brazil          106           8   
175                                       Big Fish          108           1   
176                                   True Romance          112           1   
177                                     Persepolis          114           7   

    content_rating  title_year  imdb_score      gross  duration  
0                R      
```

```python
df_directors
```

Output:
```
id          director_name
0      1         Akira Kurosawa
1      2     Alejandro Amenabar
2      3  Alejandro G. Inarritu
3      3  Alejandro G. Inarritu
4      4       Alfred Hitchcock
..   ...                    ...
173  113         Victor Fleming
174  114      Vincent Paronnaud
175  115           Wes Anderson
176  116      Wolfgang Petersen
177  117            Woody Allen

[178 rows x 2 columns]
```

```python
df_countries
```

Output:
```
id       country
0    1           USA
1    2         Italy
2    3   New Zealand
3    4         Japan
4    5        Brazil
5    6       Germany
6    7        France
7    8            UK
8    9   South Korea
9   10         India
10  11          Iran
11  12  West Germany
12  13       Denmark
13  14       Ireland
14  15        Canada
15  16     Argentina
16  17         Spain
17  18        Mexico
18  19     Australia
```

```python
df1 = pd.merge(left=df,right=df_directors, how="inner", left_on="director_id", right_on="id")
```

```python
df1.head()
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1  The Shawshank Redemption           34           1              R   
2            The Green Mile           34           1              R   
3            The Green Mile           34           1              R   
4             The Godfather           33           1              R   

   title_year  imdb_score      gross  duration  id         director_name  
0        1994         9.3   28341469       142  34        Frank Darabont  
1        1994         9.3   28341469       142  34        Frank Darabont  
2        1999         8.5  136801374       189  34        Frank Darabont  
3        1999         8.5  136801374       189  34        Frank Darabont  
4        1972         9.2  134821952       175  33  Francis Ford Coppola
```

```python
df2 = pd.merge(left=df1,right=df_countries, how="inner", left_on="country_id", right_on="id")
```

```python
df2.head()
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1  The Shawshank Redemption           34           1              R   
2            The Green Mile           34           1              R   
3            The Green Mile           34           1              R   
4             The Godfather           33           1              R   

   title_year  imdb_score      gross  duration  id_x         director_name  \
0        1994         9.3   28341469       142    34        Frank Darabont   
1        1994         9.3   28341469       142    34        Frank Darabont   
2        1999         8.5  136801374       189    34        Frank Darabont   
3        1999         8.5  136801374       189    34        Frank Darabont   
4        1972         9.2  134821952       175    33  Francis Ford Coppola   

   id_y country  
0     1     USA  
1     1     USA  
2     1     USA  
3     1     USA  
4     1     USA
```

```python
df2.shape
```

Output:
```
(408, 12)
```

<!-- grade_id: q1_test, points: 2, grade: true, solution: false -->
```python
assert_equal(df.shape, sol.df.shape)
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
""" Q2: 
Save the first ten rows of movie titles in a variable called first10, then print it
"""

# your code here

df2.head(10)
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1  The Shawshank Redemption           34           1              R   
2            The Green Mile           34           1              R   
3            The Green Mile           34           1              R   
4             The Godfather           33           1              R   
5             The Godfather           33           1              R   
6             The Godfather           33           1              R   
7    The Godfather: Part II           33           1              R   
8    The Godfather: Part II           33           1              R   
9    The Godfather: Part II           33           1              R   

   title_year  imdb_score      gross  duration  id_x         director_name  \
0        1994         9.3   28341469       142    34        Frank Darabont   
1        1994         9.3   28341469       142    34        Frank Darabont   

```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(first10, sol.first10)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
""" Q3: 
There's an extra character at the end of each movie title. 
Remove it from the data using str.replace.
And print the first ten rows of movie titles again. 
"""

# your code here

df2.head(10)
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1  The Shawshank Redemption           34           1              R   
2            The Green Mile           34           1              R   
3            The Green Mile           34           1              R   
4             The Godfather           33           1              R   
5             The Godfather           33           1              R   
6             The Godfather           33           1              R   
7    The Godfather: Part II           33           1              R   
8    The Godfather: Part II           33           1              R   
9    The Godfather: Part II           33           1              R   

   title_year  imdb_score      gross  duration  id_x         director_name  \
0        1994         9.3   28341469       142    34        Frank Darabont   
1        1994         9.3   28341469       142    34        Frank Darabont   

```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(df, sol.df)
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
""" Q4:
Who is the director with the most movies? First get the number of movies per "director_name", then save the director's name
and count as a series of length 1 called "director_with_most"
"""

# your code here

df2.groupby("director_name").count()

```

Output:
```
movie_title  director_id  country_id  content_rating  \
director_name                                                                 
Akira Kurosawa                   1            1           1               1   
Alejandro Amenabar               1            1           1               1   
Alejandro G. Inarritu            4            4           4               4   
Alfred Hitchcock                 1            1           1               1   
Andrew Stanton                   4            4           4               4   
...                            ...          ...         ...             ...   
Victor Fleming                   4            4           4               4   
Vincent Paronnaud                1            1           1               1   
Wes Anderson                     1            1           1               1   
Wolfgang Petersen                1            1           1               1   
Woody Allen                      1            1           1               1
```

```python
director_with_most = pd.DataFrame(df2.groupby("director_name")["director_id"].count())
```

```python
director_with_most.columns = ["Count"]
```

```python
director_with_most
```

Output:
```
Count
director_name               
Akira Kurosawa             1
Alejandro Amenabar         1
Alejandro G. Inarritu      4
Alfred Hitchcock           1
Andrew Stanton             4
...                      ...
Victor Fleming             4
Vincent Paronnaud          1
Wes Anderson               1
Wolfgang Petersen          1
Woody Allen                1

[117 rows x 1 columns]
```

```python
director_with_most.sort_values(by="Count")
```

Output:
```
Count
director_name            
Akira Kurosawa          1
Michel Gondry           1
Mel Gibson              1
Martin McDonagh         1
M. Night Shyamalan      1
...                   ...
Ridley Scott           16
Steven Spielberg       25
Quentin Tarantino      36
Martin Scorsese        36
Christopher Nolan      49

[117 rows x 1 columns]
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(director_with_most, sol.director_with_most)
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""Q5:
Save all of this director's movies and their ratings in a variable called all_movies_ratings, then print this variable.
(The director with the most movies you got from the last question.)
"""

# your code here

all_movies_ratings = df2[df2["director_name"] == "Christopher Nolan"]
```

```python
all_movies_ratings
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
13        The Dark Knight           16           1          PG-13        2008   
14        The Dark Knight           16           1          PG-13        2008   
15        The Dark Knight           16           1          PG-13        2008   
16        The Dark Knight           16           1          PG-13        2008   
17        The Dark Knight           16           1          PG-13        2008   
18        The Dark Knight           16           1          PG-13        2008   
19        The Dark Knight           16           1          PG-13        2008   
20              Inception           16           1          PG-13        2010   
21              Inception           16           1          PG-13        2010   
22              Inception           16           1          PG-13        2010   
23              Inception           16           1          PG-13        2010   
24              Inception           16    
```

```python
all_movies_ratings.groupby("movie_title")["imdb_score"].mean()
```

Output:
```
movie_title
Batman Begins            8.3
Inception                8.8
Interstellar             8.6
Memento                  8.5
The Dark Knight          9.0
The Dark Knight Rises    8.5
The Prestige             8.5
Name: imdb_score, dtype: float64
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(all_movies_ratings, sol.all_movies_ratings)
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""Q6:
Recommend a **random** movie that has a rating of over 8.3. 
What is the title and imdb_score of your recommendation?

Name your variables as follows:
-----------------------------------------------------------------------------
  goodmovie       <- Those movies with a rating over 8.3
  rand_int        <- The random integer index location of your recommendation
  rand_goodmovie  <- The random recommendation
"""

# Do not modify this part, it's needed for grading
import random
random.seed(0)



# your code here


```

```python
df3 = df2[df2["imdb_score"] > 8.30]
```

```python
df3.head()
```

Output:
```
movie_title  director_id  country_id content_rating  \
0  The Shawshank Redemption           34           1              R   
1  The Shawshank Redemption           34           1              R   
2            The Green Mile           34           1              R   
3            The Green Mile           34           1              R   
4             The Godfather           33           1              R   

   title_year  imdb_score      gross  duration  id_x         director_name  \
0        1994         9.3   28341469       142    34        Frank Darabont   
1        1994         9.3   28341469       142    34        Frank Darabont   
2        1999         8.5  136801374       189    34        Frank Darabont   
3        1999         8.5  136801374       189    34        Frank Darabont   
4        1972         9.2  134821952       175    33  Francis Ford Coppola   

   id_y country  
0     1     USA  
1     1     USA  
2     1     USA  
3     1     USA  
4     1     USA
```

```python
rand_goodmovie = df3.sample()
```

```python
rand_goodmovie
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
54  The Dark Knight Rises           16           1          PG-13        2012   

    imdb_score      gross  duration  id_x      director_name  id_y country  
54         8.5  448130642       164    16  Christopher Nolan     1     USA
```

<!-- grade_id: q6_test, points: 3, grade: true, solution: false -->
```python
assert_series_equal(goodmovie, sol.goodmovie)
assert_equal(rand_int, sol.rand_int)
assert_frame_equal(rand_goodmovie, sol.rand_goodmovie)
```