## Module 4

#### In this assignment, you will continue working on the movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

We have loaded and joined the data as "df" for you. Follow the instructions and finish the rest part.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: init_user, points: 0, grade: true, solution: false -->
```python
# Loading the data
import pandas as pd

xls = pd.ExcelFile('imdb.xlsx')
df = xls.parse('imdb')
df_directors = xls.parse('directors')
df_countries = xls.parse('countries')

df = pd.merge(left=df, right=df_countries, 
              how='inner', left_on='country_id', 
              right_on='id')

df = pd.merge(left=df, right=df_directors, 
              how='inner', left_on='director_id', 
              right_on='id')

print("Finished.")
```

Output:
```
Finished.
```

```python
df
```

Output:
```
movie_title  director_id  country_id content_rating  \
0    The Shawshank Redemption           34           1              R   
1    The Shawshank Redemption           34           1              R   
2              The Green Mile           34           1              R   
3              The Green Mile           34           1              R   
4               The Godfather           33           1              R   
..                        ...          ...         ...            ...   
403        Dancer in the Dark           66          13              R   
404           Pan's Labyrinth           42          17              R   
405            The Sea Inside            2          17          PG-13   
406  The Secret in Their Eyes           61          16              R   
407        Mad Max: Fury Road           39          19              R   

     title_year  imdb_score      gross  duration  id_x    country  id_y  \
0          1994         9.3   28341469       142     1        USA 
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Get the summary statistics for imdb_score and gross, then use the describe() function to summarize this visually. Save the
result in a variable called score_gross_description and print it.
"""

# your code here

df.describe()
```

Output:
```
director_id  country_id   title_year  imdb_score         gross  \
count   408.000000  408.000000   408.000000  408.000000  4.080000e+02   
mean     61.470588    2.492647  1997.247549    8.352696  1.231909e+08   
std      33.086501    3.371356    15.036565    0.287397  1.304926e+08   
min       1.000000    1.000000  1927.000000    8.000000  8.060000e+03   
25%      26.000000    1.000000  1991.000000    8.100000  2.488263e+07   
50%      71.000000    1.000000  2001.000000    8.300000  7.493675e+07   
75%      88.000000    1.000000  2008.000000    8.500000  1.879914e+08   
max     117.000000   19.000000  2016.000000    9.300000  6.232795e+08   

         duration        id_x        id_y  
count  408.000000  408.000000  408.000000  
mean   140.122549    2.492647   61.470588  
std     35.831913    3.371356   33.086501  
min     73.000000    1.000000    1.000000  
25%    116.000000    1.000000   26.000000  
50%    134.000000    1.000000   71.000000  
75%    154.500000    1.000000   88.000000
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(score_gross_description, sol.score_gross_description)
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""Q2:
What is the average rating of the director Christopher Nolan's movies? Save this value in a variable called nolan_mean and 
print.
"""

df["director_name"].value_counts()


```

Output:
```
Christopher Nolan      49
Martin Scorsese        36
Quentin Tarantino      36
Steven Spielberg       25
Ridley Scott           16
                       ..
Alejandro Amenabar      1
Danny Boyle             1
Oliver Hirschbiegel     1
James McTeigue          1
Asghar Farhadi          1
Name: director_name, Length: 117, dtype: int64
```

```python
chrisnolan = df[df["director_name"] == "Christopher Nolan"]
```

```python
chrisnolan.describe()
```

Output:
```
director_id  country_id   title_year  imdb_score         gross  \
count         49.0        49.0    49.000000   49.000000  4.900000e+01   
mean          16.0         1.0  2007.857143    8.600000  2.494235e+08   
std            0.0         0.0     4.401704    0.216025  1.776202e+08   
min           16.0         1.0  2000.000000    8.300000  2.553088e+07   
25%           16.0         1.0  2005.000000    8.500000  5.308274e+07   
50%           16.0         1.0  2008.000000    8.500000  2.053438e+08   
75%           16.0         1.0  2012.000000    8.800000  4.481306e+08   
max           16.0         1.0  2014.000000    9.000000  5.333161e+08   

         duration  id_x  id_y  
count   49.000000  49.0  49.0  
mean   143.428571   1.0  16.0  
std     19.153981   0.0   0.0  
min    113.000000   1.0  16.0  
25%    128.000000   1.0  16.0  
50%    148.000000   1.0  16.0  
75%    164.000000   1.0  16.0  
max    169.000000   1.0  16.0
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(nolan_mean, sol.nolan_mean)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""Q3: 
Create a series called 'directors' that contains each director's name and his or her average rating.  Print out the type of your variable.
Use the 'directors' series to find the average rating for Steven Spielberg.  Print the value.
"""

# your code here

df.groupby("director_name").mean()
```

Output:
```
director_id  country_id  title_year  imdb_score  \
director_name                                                            
Akira Kurosawa                 1.0         4.0      1954.0        8.70   
Alejandro Amenabar             2.0        17.0      2004.0        8.10   
Alejandro G. Inarritu          3.0         9.5      2007.5        8.10   
Alfred Hitchcock               4.0         1.0      1960.0        8.50   
Andrew Stanton                 5.0         1.0      2005.5        8.30   
...                            ...         ...         ...         ...   
Victor Fleming               113.0         1.0      1939.0        8.15   
Vincent Paronnaud            114.0         7.0      2007.0        8.00   
Wes Anderson                 115.0         1.0      2014.0        8.10   
Wolfgang Petersen            116.0        12.0      1981.0        8.40   
Woody Allen                  117.0         1.0      1977.0        8.10   

                             gross  duration  id_x   id_y  

```

```python
directors = pd.DataFrame(df.groupby("director_name")["imdb_score"].mean())
```

```python
directors
```

Output:
```
imdb_score
director_name                    
Akira Kurosawa               8.70
Alejandro Amenabar           8.10
Alejandro G. Inarritu        8.10
Alfred Hitchcock             8.50
Andrew Stanton               8.30
...                           ...
Victor Fleming               8.15
Vincent Paronnaud            8.00
Wes Anderson                 8.10
Wolfgang Petersen            8.40
Woody Allen                  8.10

[117 rows x 1 columns]
```

```python
directors.loc["Steven Spielberg"]
```

Output:
```
imdb_score    8.48
Name: Steven Spielberg, dtype: float64
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(directors, sol.directors)
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
"""Q4:
Select the non-USA (country_id=1) movies and movies made after 1960 by Hayao Miyazaki (director_id=46).
What are the years returned? Save them in a series called 'miyazaki', then print it.
"""

# your code here

hayaomiyazaki = df[df["director_name"] == "Hayao Miyazaki"]
```

```python
hayaomiyazaki
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
346         Spirited Away           46           4             PG        2001   
347         Spirited Away           46           4             PG        2001   
348         Spirited Away           46           4             PG        2001   
349     Princess Mononoke           46           4          PG-13        1997   
350     Princess Mononoke           46           4          PG-13        1997   
351     Princess Mononoke           46           4          PG-13        1997   
352  Howl's Moving Castle           46           4             PG        2004   
353  Howl's Moving Castle           46           4             PG        2004   
354  Howl's Moving Castle           46           4             PG        2004   

     imdb_score     gross  duration  id_x country  id_y   director_name  
346         8.6  10049886       125     4   Japan    46  Hayao Miyazaki  
347         8.6  10049886       125     4   Japan    46
```

```python
hayaomiyazaki["title_year"].value_counts()
```

Output:
```
1997    3
2004    3
2001    3
Name: title_year, dtype: int64
```

<!-- grade_id: q4_test, points: 2, grade: true, solution: false -->
```python
assert_frame_equal(miyazaki, sol.miyazaki)
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""Q5: 
Create a Pivot Table that shows the median rating for each director, grouped by their respective countries. Name your variable
'pivot_agg'
"""

# your code here

pivot_agg = pd.pivot_table(data=df, index="country", values="imdb_score", columns="director_name")
```

```python
pivot_agg
```

Output:
```
director_name  Akira Kurosawa  Alejandro Amenabar  Alejandro G. Inarritu  \
country                                                                    
Argentina                 NaN                 NaN                    NaN   
Australia                 NaN                 NaN                    NaN   
Brazil                    NaN                 NaN                    NaN   
Canada                    NaN                 NaN                    NaN   
Denmark                   NaN                 NaN                    NaN   
France                    NaN                 NaN                    NaN   
Germany                   NaN                 NaN                    NaN   
India                     NaN                 NaN                    NaN   
Iran                      NaN                 NaN                    NaN   
Ireland                   NaN                 NaN                    NaN   
Italy                     NaN                 NaN                    NaN   
Japan       
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(pivot_agg, sol.pivot_agg)
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""Q6:
How long did the movie Gladiator aim to keep your attention? Save the series with this information
in a variable called 'gladiator_duration', then print it.
"""

# your code here

df[df["movie_title"] == "Gladiator"]
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
203   Gladiator           90           1              R        2000   
204   Gladiator           90           1              R        2000   
205   Gladiator           90           1              R        2000   
206   Gladiator           90           1              R        2000   

     imdb_score      gross  duration  id_x country  id_y director_name  
203         8.5  187670866       171     1     USA    90  Ridley Scott  
204         8.5  187670866       171     1     USA    90  Ridley Scott  
205         8.5  187670866       171     1     USA    90  Ridley Scott  
206         8.5  187670866       171     1     USA    90  Ridley Scott
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(gladiator_duration, sol.gladiator_duration)
```