## Module 4

#### In this assignment, you will continue working on the movie data from IMDB.
- The data includes movies and ratings from the IMDB website
- Data File(s): imdb.xlsx

#### Data file contains 3 sheets:
- “imdb”: contains records of movies and ratings scraped from IMDB website
- “countries”: contains the country (of origin) names
- “directors”: contains the director names

We have loaded and joined the data as "df" for you. Follow the instructions and finish the rest part.

<!-- grade_id: init_test, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("imdb.xlsx")
from nose.tools import assert_equal
from pandas.util.testing import assert_frame_equal, assert_series_equal
```

<!-- grade_id: init_user, points: 0, grade: true, solution: false -->
```python
# Loading the data
import pandas as pd

xls = pd.ExcelFile('imdb.xlsx')
df = xls.parse('imdb')
df_directors = xls.parse('directors')
df_countries = xls.parse('countries')

df = pd.merge(left=df, right=df_countries, 
              how='inner', left_on='country_id', 
              right_on='id')

df = pd.merge(left=df, right=df_directors, 
              how='inner', left_on='director_id', 
              right_on='id')

print("Finished.")
```

Output:
```
Finished.
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
""" Q1: 
Get the summary statistics for imdb_score and gross, then use the describe() function to summarize this visually. Save the
result in a variable called score_gross_description and print it.
"""

# your code here

score_gross = ["imdb_score", "gross"]
df_score_gross = df[score_gross]
score_gross_description = df_score_gross.describe()
print(score_gross_description)
```

Output:
```
imdb_score         gross
count  178.000000  1.780000e+02
mean     8.294382  1.030402e+08
std      0.266960  1.242549e+08
min      8.000000  8.060000e+03
25%      8.100000  1.318510e+07
50%      8.200000  5.194371e+07
75%      8.475000  1.522436e+08
max      9.300000  6.232795e+08
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(score_gross_description, sol.score_gross_description)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
"""Q2:
What is the average rating of the director Christopher Nolan's movies? Save this value in a variable called nolan_mean and 
print.
"""

# your code here

chritopher_nolan = df["director_name"] == "Christopher Nolan"
nolan_mean = df[chritopher_nolan]["imdb_score"].mean()
print(nolan_mean)
```

Output:
```
8.6
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(nolan_mean, sol.nolan_mean)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
"""Q3: 
Create a series called 'directors' that contains each director's name and his or her average rating.  Print out the type of your variable.
Use the 'directors' series to find the average rating for Steven Spielberg.  Print the value.
"""

# your code here

import numpy as np
directors = df.groupby(["director_name"]).mean()["imdb_score"]
print(type(directors))
print(directors["director_name" == "Steven Spielberg"])
```

Output:
```
<class 'pandas.core.series.Series'>
8.7
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(directors, sol.directors)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
"""Q4:
Select the non-USA movies made after 1960 by Hayao Miyazaki.
Save the result in a DataFrame called 'miyazaki', then print it.

Here are the steps:
1. Query the data ('df' DataFrame) based on the following conditions:
- Non-USA movies (country_id != 1)
- Movies made after 1960 (title_year > 1960)
- Movies made by director Hayao Miyazaki (director_id == 46)
2. Save the filtered data in a DataFrame called 'miyazaki' and print it

"""

# your code here

hayao = df["director_id"] == 46
non_usa = df["country_id"] != 1
made_year = df["title_year"] > 1960

miyazaki = df[hayao & non_usa & made_year]
print(miyazaki)
```

Output:
```
movie_title  director_id  country_id content_rating  title_year  \
152         Spirited Away           46           4             PG        2001   
153     Princess Mononoke           46           4          PG-13        1997   
154  Howl's Moving Castle           46           4             PG        2004   

     imdb_score     gross  duration  id_x country  id_y   director_name  
152         8.6  10049886       125     4   Japan    46  Hayao Miyazaki  
153         8.4   2298191       134     4   Japan    46  Hayao Miyazaki  
154         8.2   4710455       119     4   Japan    46  Hayao Miyazaki
```

<!-- grade_id: q4_test, points: 2, grade: true, solution: false -->
```python
assert_frame_equal(miyazaki, sol.miyazaki)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
"""Q5: 
Create a Pivot Table that shows the median rating for each director, grouped by their respective countries. Name your variable
'pivot_agg'
"""

# your code here

pivot_agg = pd.pivot_table(
    df, index=["country", "director_name"],
    values=["imdb_score"],
    aggfunc=[np.median])
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert_frame_equal(pivot_agg, sol.pivot_agg)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
"""Q6:
How long did the movie Gladiator aim to keep your attention? Save the series with this information
in a variable called 'gladiator_duration', then print it.
"""

# your code here

df_gladiator = df[df["movie_title"] == "Gladiator"]
gladiator_duration = df_gladiator["duration"]
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
assert_series_equal(gladiator_duration, sol.gladiator_duration)
print("Success!")
```

Output:
```
Success!
```

```python

```