# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
'''
1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

'''

import csv
filepath = "ufo-sightings.csv"
ufosightings = [] 

# your code here
with open(filepath, "r") as csvfile:
    reader = csv.DictReader(csvfile)
    print(reader.fieldnames)
    for row in reader:
        ufosightings.append(row)
```

Output:
```
['datetime', 'city', 'state', 'country', 'shape', 'duration (seconds)', 'duration (hours/min)', 'comments', 'date posted', 'latitude', 'longitude ']
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
'''
2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
'''
# your code here
ufosightings_count = len(ufosightings)
print(ufosightings_count)
```

Output:
```
80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
'''
3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

Hint: Check for ufo sightings where the country is 'us'.

'''

# your code here
sightings_us = [row for row in ufosightings if row["country"] == "us"]
print(sightings_us)
```

Output:
```
IOPub data rate exceeded.
The notebook server will temporarily stop sending output
to the client in order to avoid crashing it.
To change this limit, set the config variable
`--NotebookApp.iopub_data_rate_limit`.

Current values:
NotebookApp.iopub_data_rate_limit=1000000.0 (bytes/sec)
NotebookApp.rate_limit_window=3.0 (secs)
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
'''
4. Let's find the "fireball" sighting(s) that lasted more than ten seconds in US. 
Print the the datetime and state of each.  Put the data in "fball" and print the result.

Note: Consider only the US sightings stored in "sightings_us".

- Cast the duration in seconds to a float (decimal). 
- Check if duration is greater than 10. 
- Check if the shape is "fireball".

'''

#First, define a Python function that checks if a given duration (seconds) is "valid"
def is_valid_duration(duration_as_string):
    try:
        duration = float(duration_as_string)
    except ValueError:
        return False
    else:
        return duration > 10
# your code here
fball = [row for row in sightings_us if is_valid_duration(row["duration (seconds)"]) 
         and row["shape"] == "fireball"]
for fire in fball:
    print(fire["datetime"], fire["state"])


```

Output:
```
10/10/1968 19:00 nc
10/10/1998 03:30 az
10/10/1998 23:50 tn
10/10/1999 11:00 ca
10/10/2011 20:00 ct
10/10/2011 21:00 oh
10/10/2012 20:30 nm
10/11/1987 23:59 mo
10/11/2004 22:29 tx
10/11/2006 19:00 ca
10/11/2008 22:20 nm
10/11/2012 19:30 pa
10/11/2012 20:30 nm
10/11/2013 16:35 mi
10/11/2013 20:00 oh
10/11/2013 20:15 oh
10/11/2013 20:15 ky
10/11/2013 20:20 ma
10/1/1980 20:00 al
10/1/1982 22:10 mi
10/1/1987 21:00 nh
10/1/1996 22:30 tx
10/1/1998 04:30 nm
10/1/1999 00:30 tx
10/1/2006 19:45 nv
10/1/2006 22:00 nm
10/1/2006 22:18 az
10/1/2006 23:15 co
10/1/2006 23:20 co
10/1/2011 21:43 mo
10/1/2011 21:43 mo
10/1/2012 19:30 il
10/1/2012 20:00 ca
10/1/2012 23:00 in
10/1/2013 21:30 ga
10/1/2013 22:45 ny
10/12/1998 19:00 nc
10/12/1999 05:11 or
10/12/1999 12:07 pa
10/12/1999 22:00 az
10/12/2000 19:30 ok
10/12/2003 13:08 pa
10/12/2004 21:10 wa
10/12/2005 00:00 ca
10/12/2012 00:00 fl
10/12/2012 20:30 fl
10/12/2013 07:45 mo
10/12/2013 21:00 oh
10/12/2013 21:15 oh
10/12/2013 23:28 ca
10/12/2013 24:00 o
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
'''
5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
Put the sorted list in "fballsorted" and print the result.

- Cast the duration in seconds to a float (decimal). 
- Sort in reverse order.

'''

# your code here
fballsorted = sorted(fball, key = lambda x: float(x["duration (seconds)"]), reverse = True)
fball_max= max(fball, key = lambda x: float(x["duration (seconds)"]))
               
print(fball_max["datetime"], fball_max["duration (seconds)"])
```

Output:
```
8/3/2008 21:00 10526400
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
'''
6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

- Check if the shape is "fireball".
- Cast the duration in seconds to a float (decimal).
- Get the record with the largest (max) duration in seconds.
- Get the state for the record.

'''

# your code here
state = fball_max["state"]
print(state)
```

Output:
```
va
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
'''
7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
Write code to filter out these extraneous records and get the shortest sighting overall now.  
Put the minimum duration in "min_duration" and print the result.  
Use ufosightings
Note: Consider all sightings stored in "ufosightings".

'''

# your code here
min_record = min(ufosightings, key = lambda x: x["duration (seconds)"])
min_duration = float(min_record["duration (seconds)"])
print(min_duration)
```

Output:
```
0.001
```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
'''
8. What are the top 3 shapes sighted, and how many sightings were there for each? 

Note: Consider all sightings stored in "ufosightings".

- Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
- Create a new dictionary "count" with values of that column as keys and the counts as values.
- Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

Get the top 3 and store in "top3shapes".  Print the result.

'''

#Create a new list containing values from the "shape" column in ufosightings.
# your code here
sightings_shapes = []
for row in ufosightings:
    sightings_shapes.append(row["shape"])

#Create a new dictionary with values of that column as keys and the counts as values.
# your code here
count = {}
for item in sightings_shapes:
    count[item] = count.get(item, 0) + 1
        
#Get a list of the dictionary keys and values (use the items() method) and sort them in reverse order, from greatest (most sightings) to least.
#Get and print the top 3.
# your code here
count_items = count.items()
for key,value in count_items:
    count_items_sorted = sorted(count_items, key = lambda x: x[1], reverse = True)
top3shapes = []
for i in range(3):
    top3shapes.append(count_items_sorted[i])
```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
print("Success!")
```

Output:
```
Success!
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
print("Success!")
```

Output:
```
Success!
```

```python

```