# Module 1

#### In this assignment, you will work with ufo sightings data.
- The data includes various data points about individual ufo sightings
- Data File(s): ufo-sightings.csv

<!-- grade_id: init, points: 0, grade: true, solution: false -->
```python
###########################################################
### EXECUTE THIS CELL BEFORE YOU TO TEST YOUR SOLUTIONS ###
###########################################################

import imp, os, sys
sol = imp.load_compiled("solutions", "./solutions.py")
sol.get_solutions("ufo-sightings.csv")
from nose.tools import assert_equal
```

```python
import csv
import pandas as pd
import numpy as np
```

<!-- grade_id: q1, points: 0, grade: false, solution: true -->
```python
# '''
# 1. Import the csv module. Load and read the UFO sightings data set, from the ufo-sightings.csv file, 
# into a DictReader inside a with statement.  Assume the data file is in the same directory as the code. 

# Print the field names of the data set. Iterate over the reader to put the data into a list name "ufosightings".

# '''
```

```python
df = pd.read_csv("ufo-sightings.csv", low_memory=False)
```

```python
print(df.columns)
```

Output:
```
Index(['datetime', 'city', 'state', 'country', 'shape', 'duration (seconds)',
       'duration (hours/min)', 'comments', 'date posted', 'latitude',
       'longitude '],
      dtype='object')
```

```python
df.shape
```

Output:
```
(80332, 11)
```

```python
df.head()
```

Output:
```
datetime                  city state country     shape  \
0  10/10/1949 20:30            san marcos    tx      us  cylinder   
1  10/10/1949 21:00          lackland afb    tx     NaN     light   
2  10/10/1955 17:00  chester (uk/england)   NaN      gb    circle   
3  10/10/1956 21:00                  edna    tx      us    circle   
4  10/10/1960 20:00               kaneohe    hi      us     light   

   duration (seconds) duration (hours/min)  \
0              2700.0           45 minutes   
1              7200.0              1-2 hrs   
2                20.0           20 seconds   
3                20.0             1/2 hour   
4               900.0           15 minutes   

                                            comments date posted    latitude  \
0  This event took place in early fall around 194...   4/27/2004  29.8830556   
1  1949 Lackland AFB&#44 TX.  Lights racing acros...  12/16/2005    29.38421   
2  Green/Orange circular disc over Chester&#44 En...   1/21/2008        53.2   
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings, sol.ufosightings)
```

<!-- grade_id: q2, points: 0, grade: false, solution: true -->
```python
# '''
# 2. How many sightings were there in total?  Put the count in "ufosightings_count" and print the result.
# '''
# your code here
len(df)
```

Output:
```
80332
```

<!-- grade_id: q2_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(ufosightings_count, sol.ufosightings_count)
```

<!-- grade_id: q3, points: 0, grade: false, solution: true -->
```python
# '''
# 3. How many sightings were there in the US?  Put the US sightings in "sightings_us" and print.

# Hint: Check for ufo sightings where the country is 'us'.

# '''
```

```python
sightings_us = df[df["country"] == "us"]
```

```python
sightings_us
```

Output:
```
datetime        city state country     shape  \
0      10/10/1949 20:30  san marcos    tx      us  cylinder   
3      10/10/1956 21:00        edna    tx      us    circle   
4      10/10/1960 20:00     kaneohe    hi      us     light   
5      10/10/1961 19:00     bristol    tn      us    sphere   
7      10/10/1965 23:45     norwalk    ct      us      disk   
...                 ...         ...   ...     ...       ...   
80327    9/9/2013 21:15   nashville    tn      us     light   
80328    9/9/2013 22:00       boise    id      us    circle   
80329    9/9/2013 22:00        napa    ca      us     other   
80330    9/9/2013 22:20      vienna    va      us    circle   
80331    9/9/2013 23:00      edmond    ok      us     cigar   

       duration (seconds) duration (hours/min)  \
0                  2700.0           45 minutes   
3                    20.0             1/2 hour   
4                   900.0           15 minutes   
5                   300.0            5 minutes   
7       
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_us,  sol.sightings_us)
```

<!-- grade_id: q4, points: 0, grade: false, solution: true -->
```python
# '''
# 4. Let's find the "fireball" sighting(s) that lasted more than ten seconds. 
# Print the the datetime and state of each.  Put the data in "fball" and print the result.

# - Cast the duration in seconds to a float (decimal). 
# - Check if duration is greater than 10. 
# - Check if the shape is "fireball".

# '''
```

```python
sightings_us.head()
```

Output:
```
datetime        city state country     shape  duration (seconds)  \
0  10/10/1949 20:30  san marcos    tx      us  cylinder              2700.0   
3  10/10/1956 21:00        edna    tx      us    circle                20.0   
4  10/10/1960 20:00     kaneohe    hi      us     light               900.0   
5  10/10/1961 19:00     bristol    tn      us    sphere               300.0   
7  10/10/1965 23:45     norwalk    ct      us      disk              1200.0   

  duration (hours/min)                                           comments  \
0           45 minutes  This event took place in early fall around 194...   
3             1/2 hour  My older brother and twin sister were leaving ...   
4           15 minutes  AS a Marine 1st Lt. flying an FJ4B fighter/att...   
5            5 minutes  My father is now 89 my brother 52 the girl wit...   
7           20 minutes  A bright orange color changing to reddish colo...   

  date posted    latitude  longitude   
0   4/27/2004  29.8830556  -97.94
```

```python
sightings_us["shape"].value_counts()
```

Output:
```
light        13473
triangle      6549
circle        6118
fireball      5148
unknown       4567
other         4466
sphere        4347
disk          4121
oval          3032
formation     1990
cigar         1643
changing      1584
rectangle     1070
flash         1062
cylinder      1027
diamond        928
chevron        815
teardrop       592
egg            581
cone           247
cross          188
delta            6
round            2
hexagon          1
crescent         1
pyramid          1
flare            1
changed          1
Name: shape, dtype: int64
```

```python
df["duration (seconds)"] = df["duration (seconds)"].astype("float")
```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 80332 entries, 0 to 80331
Data columns (total 11 columns):
 #   Column                Non-Null Count  Dtype  
---  ------                --------------  -----  
 0   datetime              80332 non-null  object 
 1   city                  80332 non-null  object 
 2   state                 74535 non-null  object 
 3   country               70662 non-null  object 
 4   shape                 78400 non-null  object 
 5   duration (seconds)    80332 non-null  float64
 6   duration (hours/min)  80332 non-null  object 
 7   comments              80317 non-null  object 
 8   date posted           80332 non-null  object 
 9   latitude              80332 non-null  object 
 10  longitude             80332 non-null  float64
dtypes: float64(2), object(9)
memory usage: 6.7+ MB
```

```python
fireball = sightings_us[(sightings_us["shape"] == "fireball") & (sightings_us["duration (seconds)"] > 10)]
```

```python
fireball.head()
```

Output:
```
datetime                  city state country     shape  \
11   10/10/1968 19:00               brevard    nc      us  fireball   
73    10/10/1998 3:30  phoenix (north part)    az      us  fireball   
81   10/10/1998 23:50                delano    tn      us  fireball   
85   10/10/1999 11:00             san diego    ca      us  fireball   
241  10/10/2011 20:00            middletown    ct      us  fireball   

     duration (seconds) duration (hours/min)  \
11                180.0            3 minutes   
73                 15.0         10 to 15 sec   
81                 15.0           15 seconds   
85               3600.0               1 hour   
241              1200.0        15-20 minutes   

                                              comments date posted  \
11   silent red /orange mass of energy floated by t...   6/12/2008   
73   bright green moving north to north east. it wa...   2/16/1999   
81   I came home from work one night and I seen a b...   3/21/2003   
85   At 10:45 in 
```

```python
fireball.shape[0]
```

Output:
```
3730
```

<!-- grade_id: q4_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fball, sol.fball)
```

<!-- grade_id: q5, points: 0, grade: false, solution: true -->
```python
# '''
# 5. Sort the above list by duration. What was the datetime and duration of the longest sighting?  
# Put the sorted list in "fballsorted" and print the result.

# - Cast the duration in seconds to a float (decimal). 
# - Sort in reverse order.

# '''

fireball.head()

```

Output:
```
datetime                  city state country     shape  \
11   10/10/1968 19:00               brevard    nc      us  fireball   
73    10/10/1998 3:30  phoenix (north part)    az      us  fireball   
81   10/10/1998 23:50                delano    tn      us  fireball   
85   10/10/1999 11:00             san diego    ca      us  fireball   
241  10/10/2011 20:00            middletown    ct      us  fireball   

     duration (seconds) duration (hours/min)  \
11                180.0            3 minutes   
73                 15.0         10 to 15 sec   
81                 15.0           15 seconds   
85               3600.0               1 hour   
241              1200.0        15-20 minutes   

                                              comments date posted  \
11   silent red /orange mass of energy floated by t...   6/12/2008   
73   bright green moving north to north east. it wa...   2/16/1999   
81   I came home from work one night and I seen a b...   3/21/2003   
85   At 10:45 in 
```

```python
fireball.sort_values(by="duration (seconds)", ascending=False)
```

Output:
```
datetime                            city state country  \
71172    8/3/2008 21:00                  virginia beach    va      us   
23721             23998                  newtown square    pa      us   
62793     7/5/1996 0:00                        oak hill    wv      us   
21793  12/31/2013 21:54                            mesa    az      us   
21857   12/3/1999 19:30                            lodi    ca      us   
...                 ...                             ...   ...     ...   
34593    3/3/2013 21:57                   bullhead city    az      us   
45406     5/3/2013 2:30                          venice    ca      us   
26898   2/16/2003 21:30                            arco    id      us   
5698   10/31/2000 18:45                 rochester hills    mi      us   
9785   11/16/1999 19:00  cincinnati (eastern outskirts)    oh      us   

          shape  duration (seconds) duration (hours/min)  \
71172  fireball          10526400.0             4 months   
23721  fireball   
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(fballsorted, sol.fballsorted)
```

<!-- grade_id: q6, points: 0, grade: false, solution: true -->
```python
# '''
# 6. What state had the longest lasting "fireball"?   Put the state in "state" and print the result.

# - Check if the shape is "fireball".
# - Cast the duration in seconds to a float (decimal).
# - Get the record with the largest (max) duration in seconds.
# - Get the state for the record.

# '''
```

```python
fireball[fireball["duration (seconds)"] == fireball["duration (seconds)"].max()]
```

Output:
```
datetime            city state country     shape  \
71172  8/3/2008 21:00  virginia beach    va      us  fireball   

       duration (seconds) duration (hours/min)  \
71172          10526400.0             4 months   

                                                comments date posted  \
71172  this object was very high up and emmited no so...   8/12/2008   

         latitude  longitude   
71172  36.8527778  -75.978333
```

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(state, sol.state)
```

<!-- grade_id: q7, points: 0, grade: false, solution: true -->
```python
# '''
# 7. Let's assume that any sighting (of any shape) of 0 seconds is insignificant. 
# Write code to filter out these extraneous records and get the shortest sighting overall now.  
# Put the minimum duration in "min_duration" and print the result.

# '''
```

```python
df.head()
```

Output:
```
datetime                  city state country     shape  \
0  10/10/1949 20:30            san marcos    tx      us  cylinder   
1  10/10/1949 21:00          lackland afb    tx     NaN     light   
2  10/10/1955 17:00  chester (uk/england)   NaN      gb    circle   
3  10/10/1956 21:00                  edna    tx      us    circle   
4  10/10/1960 20:00               kaneohe    hi      us     light   

   duration (seconds) duration (hours/min)  \
0              2700.0           45 minutes   
1              7200.0              1-2 hrs   
2                20.0           20 seconds   
3                20.0             1/2 hour   
4               900.0           15 minutes   

                                            comments date posted    latitude  \
0  This event took place in early fall around 194...   4/27/2004  29.8830556   
1  1949 Lackland AFB&#44 TX.  Lights racing acros...  12/16/2005    29.38421   
2  Green/Orange circular disc over Chester&#44 En...   1/21/2008        53.2   
```

```python
df["duration (seconds)"].value_counts()
```

Output:
```
300.0        8635
120.0        7110
600.0        6312
60.0         5866
180.0        5065
             ... 
7894800.0       1
209.0           1
304.0           1
1053.0          1
4560.0          1
Name: duration (seconds), Length: 533, dtype: int64
```

```python
sightings_us.groupby(by='duration (seconds)')["country"].count()
```

Output:
```
duration (seconds)
0.01           4
0.02           1
0.05           3
0.08           2
0.10           4
              ..
7894800.00     1
9468000.00     1
10526400.00    2
52623200.00    2
66276000.00    1
Name: country, Length: 466, dtype: int64
```

<!-- grade_id: q7_test, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(min_duration, sol.min_duration)
```

<!-- grade_id: q8, points: 0, grade: false, solution: true -->
```python
# '''
# 8. What are the top 3 shapes sighted, and how many sightings were there for each? 

# - Create a new list "sightings_shapes" containing values from the "shape" column in ufosightings.  
# - Create a new dictionary "count" with values of that column as keys and the counts as values.
# - Get a list of the dictionary keys and values using the items() method.  This will return a list of key:value pairs.
# Sort the list of key:value pairs in reverse order, from greatest (most sightings) to least.

# Get the top 3 and store in "top3shapes".  Print the result.

# '''
```

```python
df["shape"].value_counts()
```

Output:
```
light        16565
triangle      7865
circle        7608
fireball      6208
other         5649
unknown       5584
sphere        5387
disk          5213
oval          3733
formation     2457
cigar         2057
changing      1962
flash         1328
rectangle     1297
cylinder      1283
diamond       1178
chevron        952
egg            759
teardrop       750
cone           316
cross          233
delta            7
crescent         2
round            2
dome             1
hexagon          1
pyramid          1
flare            1
changed          1
Name: shape, dtype: int64
```

```python
sighting_shapes = pd.DataFrame(df["shape"].value_counts())
```

```python
sighting_shapes
```

Output:
```
shape
light      16565
triangle    7865
circle      7608
fireball    6208
other       5649
unknown     5584
sphere      5387
disk        5213
oval        3733
formation   2457
cigar       2057
changing    1962
flash       1328
rectangle   1297
cylinder    1283
diamond     1178
chevron      952
egg          759
teardrop     750
cone         316
cross        233
delta          7
crescent       2
round          2
dome           1
hexagon        1
pyramid        1
flare          1
changed        1
```

```python
sighting_shapes.reset_index(inplace=True)
```

```python
sighting_shapes
```

Output:
```
index  shape
0       light  16565
1    triangle   7865
2      circle   7608
3    fireball   6208
4       other   5649
5     unknown   5584
6      sphere   5387
7        disk   5213
8        oval   3733
9   formation   2457
10      cigar   2057
11   changing   1962
12      flash   1328
13  rectangle   1297
14   cylinder   1283
15    diamond   1178
16    chevron    952
17        egg    759
18   teardrop    750
19       cone    316
20      cross    233
21      delta      7
22   crescent      2
23      round      2
24       dome      1
25    hexagon      1
26    pyramid      1
27      flare      1
28    changed      1
```

```python
sighting_shapes.columns = ["shape","count"]
```

```python
sighting_shapes
```

Output:
```
shape  count
0       light  16565
1    triangle   7865
2      circle   7608
3    fireball   6208
4       other   5649
5     unknown   5584
6      sphere   5387
7        disk   5213
8        oval   3733
9   formation   2457
10      cigar   2057
11   changing   1962
12      flash   1328
13  rectangle   1297
14   cylinder   1283
15    diamond   1178
16    chevron    952
17        egg    759
18   teardrop    750
19       cone    316
20      cross    233
21      delta      7
22   crescent      2
23      round      2
24       dome      1
25    hexagon      1
26    pyramid      1
27      flare      1
28    changed      1
```

```python
sighting_shapes.head(3)
```

Output:
```
shape  count
0     light  16565
1  triangle   7865
2    circle   7608
```

<!-- grade_id: q8_test_a, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(sightings_shapes, sol.sightings_shapes)
```

<!-- grade_id: q8_test_b, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(count, sol.count)
```

<!-- grade_id: q8_test_c, points: 1, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

assert_equal(top3shapes, sol.top3shapes)
```